---
title: "Introducci칩n a R"
author: "Fabi치n Gil"
date: "2024-11-01"
output:
  learnr::tutorial:
    progressive: true
    allow_skip: true
    df_print: default
runtime: shiny_prerendered
description: >
  Introducci칩n al an치lisis estad칤stico usando R, 
  estad칤stica descriptiva, tablas de frecuencia y gr치ficos.
editor_options: 
  markdown: 
    wrap: 72
---

```{r setup, include=FALSE}
# Attach packages:

library(learnr)
library(fontawesome)
library(ggplot2)
gradethis::gradethis_setup()
tutorial_options(exercise.timelimit = 60)

# Global code chunk settings:
knitr::opts_chunk$set(echo = FALSE)

# data
load("data/nac.rdata")
```

## 1. Introducci칩n

```{r, echo=FALSE, out.width="100%", fig.align = "center"}
knitr::include_graphics("images/banner.png")  
```

<a rel="license" href="http://creativecommons.org/licenses/by-nc-sa/4.0/"><img src="https://i.creativecommons.org/l/by-nc-sa/4.0/88x31.png" alt="Licencia de Creative Commons" style="border-width:0"/></a><br />Este
obra est치 bajo una
<a rel="license" href="http://creativecommons.org/licenses/by-nc-sa/4.0/">licencia
de Creative Commons Reconocimiento-NoComercial-CompartirIgual 4.0
Internacional</a>.

Este taller tiene como objetivo introducir el an치lisis de datos
empleando el programa estad칤stico [R](https://cran.r-project.org/).

Los temas que ser치n cubiertos en este taller son:

1.  C치lculo de estad칤sticas descriptivas: tendencia central y
    dispersi칩n.
2.  C치lculo de frecuencias absolutas y relativas, tablas de frecuencias.
3.  Construcci칩n de gr치ficos estad칤sticos

### 

### 쮺칩mo desarrollar este tutorial?

A lo largo del tutorial, encontrar치 ejemplos de c칩digo, ejercicios de
programaci칩n y preguntas de an치lisis. Los ejemplos de c칩digo y los
ejercicios pretenden introducir habilidades b치sicas de R. No es
necesario tener conocimientos previos sobre R para el desarrollo de este
tutorial.

#### `r fa("fas fa-robot", fill = "orange")` = Ejemplo de c칩digo

El icono del robot identifica los ejemplos de c칩digo a lo largo del
tutorial. Estos ejemplos incluyen c칩digo preescrito y funcional que
puede explorar y ejecutar al presionar el bot칩n 'Run code' en la esquina
superior derecha.

```{r, echo=FALSE, out.width="100%", fig.align = "center"}
knitr::include_graphics("images/ejemplo.png")  
```

Puede explorar y modificar libremente el c칩digo del ejemplo, **no hay
forma de da침ar el c칩digo de ejemplo**. Si realiza cambios en alguno de
los ejemplos y recibe un mensaje de error, simplemente haga clic en el
bot칩n 'Start Over' en la esquina superior izquierda, y el c칩digo se
restablecer치 autom치ticamente a su configuraci칩n predeterminada.

#### `r fa("fas fa-keyboard", fill = "green")` = Ejercicio de c칩digo

El icono del teclado se침ala ejercicios de pr치ctica en los que deber치
ingresar el c칩digo manualmente para obtener los resultados que
solucionen las preguntas planteadas. Si siente que se ha quedado sin
ideas o no sabes c칩mo continuar con el ejercicio, simplemente haga clic
en el bot칩n 'Solution' o 'Hint'.

```{r, echo=FALSE, out.width="100%", fig.align = "center"}
knitr::include_graphics("images/ejercicio.png")  
```

#### `r fa("brain", fill = "#2c5697")` = Preguntas de an치lisis

El icono del cerebro se침ala preguntas que debes reflexionar y responder.
En algunos casos, no existe una respuesta correcta o incorrecta; en
otros casos puede haber m칰ltiples respuestas v치lidas. Muchas de estas
preguntas de reflexi칩n estar치n relacionadas con un ejemplo o ejercicio
de c칩digo. En esos casos, aseg칰rese de ejecutar el c칩digo
correspondiente para poder ver la salida necesaria que le ayudar치 a
responder las preguntas.

**NOTE:** Si desea profundizar en el manejo del programa R y planea
emplearlo de forma frecuente se recomienda instalarlo en su computador
游눹, esto le permitir치 trabajar sin conexi칩n a la Internet y emplear un
mayor conjunto de herramientas. Para m치s informaci칩n consulte este
[material](https://github.com/introbios/R-Rstudio). Tambi칠n puede
consultar en el Departamento de Epidemiolog칤a Cl칤nica y Bioestad칤stica
por el ***Semillero de bioestad칤stica***

## 2. Problema de investigaci칩n

```{r, echo=FALSE, out.width="50%", fig.align = "center"}
knitr::include_graphics("images/bp.png")  
```

El bajo peso al nacer, que se define como un peso inferior a 2500 gramos
al momento del parto, es un importante indicador de la salud neonatal.
Este fen칩meno est치 asociado con un aumento en las tasas de morbilidad y
mortalidad infantil, as칤 como con complicaciones a largo plazo en el
desarrollo f칤sico y cognitivo de los ni침os. Diversos factores
contribuyen al BPN, incluyendo la nutrici칩n materna, la edad de la
madre, el consumo de tabaco y el acceso a atenci칩n prenatal adecuada.

Realizar un an치lisis estad칤stico sobre la prevalencia y los
determinantes del bajo peso al nacer es fundamental para identificar
poblaciones en riesgo y comprender mejor las din치micas que lo afectan.
Este tipo de an치lisis puede ayudar a desarrollar intervenciones
efectivas que mejoren los resultados de salud materno-infantil,
permitiendo a los profesionales de la salud implementar estrategias
dirigidas a reducir la incidencia del bajo peso al nacer y sus
consecuencias asociadas.Para este fin emplearemos la informaci칩n
recolectada sobre una muestra de nacimientos simples observados en el
condado de King (Washington, USA)

### 

### 2.1 Datos para el an치lisis

El archivo de datos contiene la informaci칩n de 13 variables recolectadas
para 2500 nacimientos. Las variables recogen informaci칩n de la madre y
del reci칠n nacido como se muestra en la tabla 1:

| **Variable** | **Descripci칩n**                                 | **Valores**                                                                  |
|------------------------|------------------------|------------------------|
| *sexo*       | Sexo del reci칠n nacido                          | Femenino, Masculino                                                          |
| *edadm*      | Edad de la madre al nacimiento                  | a침os                                                                         |
| *etnia*      | Grupo 칠tnico de la madre                        | Asi치tica, Afro, Hispana, Blanca, Otro                                        |
| *paridad*    | N칰mero previo de hijos                          | 0, 1, 2, ...                                                                 |
| *casada*     | Estado conyugal de la madre                     | soltera, casada                                                              |
| *pesonac*    | Peso al nacer                                   | gramos                                                                       |
| *bajopeso*   | Bajo peso al nacer (\< 2500 gramos)             | No, Si                                                                       |
| *fuma*       | 쮽um칩 la madre durante el embarazo?             | No, Si                                                                       |
| *bebe*       | 쮺onsumi칩 la madre alcohol durante el embarazo? | No, Si                                                                       |
| *pesopre*    | Peso de la madre previo al embarazo             | libras                                                                       |
| *pesogan*    | Peso ganado durante el embarazo                 | libras                                                                       |
| *educacion*  | Educaci칩n de la madre                           | a침os                                                                         |
| *edadges*    | Edad gestacional en categor칤as                  | Prematuro extremo, prematuro moderado, prematuro leve, a t칠rmino, postmaduro |

<center>**Tabla 1: variables disponibles en el archivo de datos
*nac.rdata***</center>

### Lectura del archivo de datos en R

La informaci칩n recolectada sobre los reci칠n nacidos y sus madres se
encuentra disponible en el archivo **nac.Rdata**, un archivo en formato
del programa R.

El primer paso consiste en leer los datos en R. Para leer un archivo en
formato R se emplea el comando `load()`, es importante incluir la ruta
completa del archivo en el comando.

```{r nac-load, echo=TRUE, eval=FALSE}
load("datos/nac.rdata")
```

Como resultado los datos han sido le칤dos en R y ahora se encuentran un
objeto `dataframe` de R con el nombre **nac**; este nombre es importante
para acceder y hacer referencia al conjunto de datos dentro de R.

### Descripci칩n del archivo de datos en R

Ahora podemos revisar la caracter칤sticas del archivo de datos, como el
n칰mero de observaciones, el n칰mero de variables y el tipo de datos de
cada variable.

##### `r fa("fas fa-robot", fill = "orange")` Ejemplo 2.1

Describa la estructura del archivo de datos, presione "run code" para
ver el resultado

```{r nac-desc, exercise = TRUE}
str(nac)
```

### 

<font color="#41A9FB"> En la salida se puede identificar que el archivo
de datos se compone de **2500** observaciones y **13** variables. Se
identifican los nombres de las variables, junto con el tipo de datos
correspondiente a cada una, y se presenta el listado de los primeros
valores almacenados en cada variable. Por ejemplo, la variable *sexo* es
de tipo `factor` con los niveles "Femenino" y "Masculino"; mientras que
la variable *edam* es de tipo `num칠rico` y los primeros valores
almacenados son 31, 23, 24, ... </font>

## 3. Estad칤sticas descriptivas para variables cualitativas

### 3.1 Describiendo una variable

#### 3.1.1 Frecuencias absolutas

Iniciemos describiendo el sexo de los reci칠n nacidos, variable *gender*,
la cual es una variable cualitativa en escala nominal. El inter칠s es
conocer el n칰mero de ni침os y ni침as nacidas

Para calcular las frecuencias absolutas de una variable se emplea la
funci칩n `table()`. Se debe especificar dentro del parentesis el nombre
de la variable precedida de signo \$ y el nombre del objeto que la
contiene.

##### `r fa("fas fa-robot", fill = "orange")` Ejemplo 3.1

Describa el sexo de los reci칠n nacidos, presione "run code" para ver el
resultado

```{r sex-desc, exercise = TRUE}
table(nac$sexo)
```

### 

<font color="#41A9FB"> Se observa que 1209 reci칠n nacidos fueron de sexo
femenino y 1291 fueron de sexo masculino</font>

### 

#### 3.1.2 Frecuencias relativas

Ahora el inter칠s es conocer que parte del total de reci칠n nacidos
corresponde a cada uno de los sexos.

Para obtener las proporciones de cada categor칤a de la variable se emplea
la funci칩n `prop.table()`, este comando se debe aplicar sobre la funci칩n
`table()`

##### `r fa("fas fa-robot", fill = "orange")` Ejemplo 3.2

Calcule la proporci칩n de cada sexo de los reci칠n nacidos, presione "run
code" para ver el resultado

```{r sex-desc2, exercise = TRUE}
prop.table(table(nac$sexo))
```

### 

<font color="#41A9FB"> La proporci칩n de mujeres y hombres es 0.484 y
0.516, respectivamente.</font>

### 

Estos valores pueden ser multiplicados por 100 para ser convertidos en
porcentajes y mejorar su interpretaci칩n

##### `r fa("fas fa-robot", fill = "orange")` Ejemplo 3.3

```{r sex-desc3, exercise = TRUE}
prop.table(table(nac$sexo))*100
```

### 

<font color="#41A9FB"> Este resultado indica que aproximadamente de cada
100 nacidos vivos 48 son del sexo femenino 游놌 y 52 son de sexo masculino
游녽</font>

#### 3.1.3 Ejercicios de aplicaci칩n

Ahora es el momento de aplicar los conceptos revisados

Es su turno! Complete los espacios en blanco en el c칩digo de los
siguiente escenarios para proporcionar una descripci칩n con frecuencias
absolutas o relativas de las variables indicadas. Tenga presente que
tiene la opci칩n de ver la soluci칩n (춰pero aseg칰rese de intentar
resolverla manualmente primero!)

##### `r fa("fas fa-keyboard", fill = "green")` Ejercicio 3.1

Describa el grupo 칠tnico de las madres `etnia`

```{r exercise_3_1, exercise = TRUE, exercise.blanks = "___+"}
table(nac$______)
```

```{r exercise_3_1-solution}
table(nac$etnia)
```

```{r exercise_3_1-check}
grade_code()
```

### 

##### `r fa("fas fa-keyboard", fill = "green")` Ejercicio 3.2

Describa el h치bito de fumar de las madres durante el embarazo `fuma`

```{r exercise_3_2, exercise = TRUE, exercise.blanks = "___+"}
table(______$fuma)
```

```{r exercise_3_2-solution}
table(nac$fuma)
```

```{r exercise_3_2-check}
grade_code()
```

### 

##### `r fa("fas fa-keyboard", fill = "green")` Ejercicio 3.3

Describa el h치bito de beber de las madres durante el embarazo `bebe`

```{r exercise_3_3, exercise = TRUE, exercise.blanks = "___+"}
table(______)
```

```{r exercise_3_3-solution}
table(nac$bebe)
```

```{r exercise_3_3-check}
grade_code()
```

### 

##### `r fa("fas fa-keyboard", fill = "green")` Ejercicio 3.4

쮺u치l es el porcentaje de madres en cada grupo 칠tnico `etnia`?

```{r exercise_3_4, exercise = TRUE, exercise.blanks = "___+"}
prop.table(table(nac$______))*100
```

```{r exercise_3_4-solution}
prop.table(table(nac$etnia))*100
```

```{r exercise_3_4-check}
grade_code()
```

### 

##### `r fa("fas fa-keyboard", fill = "green")` Ejercicio 3.5

쮺u치l es el porcentaje de madres que fumaron en el embarazo `fuma`?

```{r exercise_3_5, exercise = TRUE, exercise.blanks = "___+"}
______(table(______$fuma))*100
```

```{r exercise_3_5-solution}
prop.table(table(nac$fuma))*100
```

```{r exercise_3_5-check}
grade_code()
```

### 

##### `r fa("fas fa-keyboard", fill = "green")` Ejercicio 3.6

쮺u치l es el porcentaje de madres que bebieron en el embarazo `bebe`?

```{r exercise_3_6, exercise = TRUE, exercise.blanks = "___+"}
prop.table(______)*100
```

```{r exercise_3_6-solution}
prop.table(table(nac$bebe))*100
```

```{r exercise_3_6-check}
grade_code()
```

### 3.2 Describiendo dos variables de forma conjunta

#### 3.2.1 Frecuencias absolutas

El inter칠s es describir el consumo de alcohol 游꽄 y cigarrillo 游뛃 de las
madres durante el embrazo, dado que ambas variables son de tipo
cualitativo en escala nominal se puede construir una tabla de
frecuencias a dos v칤as

El comando para construir una tabla a dos v칤as es
`table(<variable filas>, <variable columnas>)`, en esta oportunidad se
debe especificar las dos variables separadas por una coma.

##### `r fa("fas fa-robot", fill = "orange")` Ejemplo 3.4

```{r ejemplo-3_4, exercise = TRUE}
table(nac$bebe, nac$fuma)
```

### 

<font color="#41A9FB"> La tabla obtenida muestra la primera variable
usada en el comando (*bebe*) en las filas y la segunda variable (*fuma*)
en las columnas. De esta forma 6 madres fumaron y consumieron alcohol en
el embarazo y 2302 madres no consumieron ninguna de las sustancias. Por
otra parte, 169 madres fumaron, pero no bebieron; mientras que 23 madres
no fumaron, pero si bebieron durante el embarazo.</font>

#### 3.2.2 Frecuencias relativas

Ahora la pregunta es, del total de madres en el estudio, 쯤u칠 porcentaje
consumieron alcohol y cigarrillo durante el embarazo?

Al igual que en las tablas de una v칤a, el comando para obtener
proporciones sobre es `prop.table()`, este comando se debe aplicar sobre
la tabla de dos v칤as. Para obtener porcentajes se debe multiplicar por
100.

##### `r fa("fas fa-robot", fill = "orange")` Ejemplo 3.5

```{r ejemplo-3_5, exercise = TRUE}
prop.table(table(nac$bebe, nac$fuma))*100
```

### 

<font color="#41A9FB"> Se Observa que el 0.2% del total de madres fum칩 y
consumi칩 alcohol durante el embrazo, mientras que el 92.1% de las madres
no consumi칩 ninguna de estas dos sustancias. </font>

### 

Sobre una tabla de dos v칤as tambi칠n es posible calcular porcentajes
sobre los totales de fila o columna. Ahora el inter칠s es conocer de
acuerdo con el consumo de alcohol durante el embarazo, 쯤u칠 porcentaje
de madres fumaron y no fumaro?; es decir se desea conocer el porcentaje
sobre las filas de la tabla.

Para obtener proporciones con respecto a las filas de una tabla se
emplea el comando `prop.table(<tabla>,1)`, el n칰mero 1 significa que las
proporciones se calculan sobre el total de cada fila.

##### `r fa("fas fa-robot", fill = "orange")` Ejemplo 3.6

```{r ejemplo-3_6, exercise = TRUE}
prop.table(table(nac$bebe, nac$fuma),1)*100
```

### 

<font color="#41A9FB"> Los resultados muestran que del grupo de madres
que no consumieron alcohol en el embarazo el 6.8% si fumaron, mientras
que del grupo de madres que bebi칩 alcohol el 20.7% tambi칠n fumaron.
</font>

### 

Una pregunta alternativa es describir el comportamiento del consumo de
alcohol de acuerdo al consumo de cigarrillo en el embarazo.

Para obtener proporciones con respecto a las columnas de una tabla se
emplea el comando `prop.table(<tabla>,2)`, el n칰mero 2 significa que las
proporciones se calculan sobre el total de cada columna

##### `r fa("fas fa-robot", fill = "orange")` Ejemplo 3.7

```{r ejemplo-3_7, exercise = TRUE}
prop.table(table(nac$bebe, nac$fuma),2)*100
```

### 

<font color="#41A9FB"> En el grupo de madres que no fumaron en el
embarazo el 1.0% tampoco consumieron alcohol, mientras que del grupo de
madres que fum칩 el 3.4% tambien consumieron alcohol. </font>

### 

#### 3.3 Ejercicios de aplicaci칩n

Ahora es el momento de aplicar los conceptos revisados

Complete las l칤neas de c칩digo en cada una de las celdas, se desea
describir con frecuencias absolutas y relativas de forma conjunta las
variables grupo 칠tnico `etnia` y estado conyugal de la madre `casada`

##### `r fa("fas fa-keyboard", fill = "green")` Ejercicio 3.7

쮺u치ntas mujeres son casadas en cada grupo 칠tnico en este estudio?

```{r exercise_3_7, exercise = TRUE, exercise.blanks = "___+"}
table(nac$______, nac$______)
```

```{r exercise_3_7-solution}
table(nac$etnia, nac$casada)
```

```{r exercise_3_7-check}
grade_code()
```

### 

##### `r fa("fas fa-keyboard", fill = "green")` Ejercicio 3.8

Del total de madres en este estudio, 쯤u칠 porcentaje son casadas e
hispanas?

```{r exercise_3_8, exercise = TRUE, exercise.blanks = "___+"}
______(______(nac$etnia, nac$casada))*______
```

```{r exercise_3_8-solution}
prop.table(table(nac$etnia, nac$casada))*100
```

```{r exercise_3_8-check}
grade_code()
```

### 

##### `r fa("fas fa-keyboard", fill = "green")` Ejercicio 3.9

De las madres no casadas, 쯤u칠 porcentaje son asi치ticas?

```{r exercise_3_9, exercise = TRUE, exercise.blanks = "___+"}
prop.table(table(nac$etnia, nac$casada),______)*100
```

```{r exercise_3_9-solution}
prop.table(table(nac$etnia, nac$casada),2)*100
```

```{r exercise_3_9-check}
grade_code()
```

### 

##### `r fa("fas fa-keyboard", fill = "green")` Ejercicio 3.10

쮺u치l es el porcentaje de madres casadas en cada grupo 칠tnico?

```{r exercise_3_10, exercise = TRUE, exercise.blanks = "___+"}
prop.table(table(nac$etnia, nac$casada),______)*100
```

```{r exercise_3_10-solution}
prop.table(table(nac$etnia, nac$casada),1)*100
```

```{r exercise_3_10-check}
grade_code()
```

## 4. Estad칤sticas descriptivas para variables cuantitativas

La descripci칩n de las variables cuantitativas puede hacerse con
referencia a estad칤sticas de tendencia central: un valor "t칤pico" o
central, estad칤sticas de dispersi칩n: respecto a la variabilidad que se
observa entre los valores de la variable con respecto a un valor; o
estad칤sticas de posici칩n: posici칩n relativa de los valores de la
variable respecto al conjunto total de datos.

### 4.1 Estad칤sticas de tendencia central

#### 4.1.1 Media o promedio ar칤tmetico

쮺u치l es el peso promedio de los reci칠n nacidos?

##### `r fa("fas fa-robot", fill = "orange")` Ejemplo 4.1

```{r ejemplo-4_1, exercise = TRUE}
mean(nac$pesonac)
```

### 

<font color="#41A9FB"> El peso al nacer de los reci칠n nacidos del
estudio estuvo alrededor de los 3414 gramos. </font>

#### 4.1.2 Mediana

쮺u치l es el peso mediano de los beb칠s al nacer?

##### `r fa("fas fa-robot", fill = "orange")` Ejemplo 4.2

```{r ejemplo-4_2, exercise = TRUE}
median(nac$pesonac)
```

### 

<font color="#41A9FB"> El 50% de los reci칠n nacidos tuvieron un peso
menor a 3445 gramos al nacer</font>

#### 4.1.3 Moda

쮺u치l es la moda del n칰mero previo de hijos?

##### `r fa("fas fa-robot", fill = "orange")` Ejemplo 4.3

```{r ejemplo-4_3, exercise = TRUE}
table(nac$paridad)
```

### 

<font color="#41A9FB"> La moda corresponde al valor de la variable que
con mayor frecuencia ocurri칩, para esto se observan las frecuencias
absolutas en una tabla de frecuencias. Se observa que con mayor
frecuencia las madres del estudio no ten칤an hijos previamente.</font>

### 4.2 Estad칤sticas de dispersi칩n

#### 4.2.1 Rango

쮺u치l es el rango del peso al nacer?

##### `r fa("fas fa-robot", fill = "orange")` Ejemplo 4.4

```{r ejemplo-4_4, exercise = TRUE}
range(nac$pesonac)
```

### 

<font color="#41A9FB"> El rango corresponde a la distancia entre los
valores m칤nimo y m치ximo observados de la variable. El rango del peso al
nacer est치 entre 255 y 5175 gramos.</font>

#### 4.2.2 Desviaci칩n est치ndar

쮺u치l es la desviaci칩n est치ndar del peso al nacer?

##### `r fa("fas fa-robot", fill = "orange")` Ejemplo 4.5

```{r ejemplo-4_5, exercise = TRUE}
sd(nac$pesonac)
```

### 

<font color="#41A9FB"> En promedio, los valores del peso al nacer se
alejan 559.4 gramos del valor promedio de la variable.</font>

#### 4.2.3 Coeficiente de variaci칩n

쮺u치l es el coeficiente de variaci칩n del peso al nacer?

##### `r fa("fas fa-robot", fill = "orange")` Ejemplo 4.6

```{r ejemplo-4_6, exercise = TRUE}
sd(nac$pesonac) / mean(nac$pesonac) * 100
```

### 

<font color="#41A9FB">El coeficiente de variaci칩n indica relativamente
la cantidad de variaci칩n de la variable con respecto al promedio, se
expresa en porcentaje. la variabilidad del peso al nacer es el 16.4% del
valor promedio.</font>

### 4.3 Estad칤sticas de posici칩n

쮺u치l es el percentil 10 del peso al nacer?

##### `r fa("fas fa-robot", fill = "orange")` Ejemplo 4.7

```{r ejemplo-4_7, exercise = TRUE}
quantile(nac$pesonac, 0.10)
```

### 

<font color="#41A9FB"> El 10% de los nacidos vivos con menores pesos se
encuentran por debajo de los 2750 gramos. Del mismo modo, el 90% de los
mayores pesos est치n por encima de los 2750 gramos al nacer.</font>

### 

쮺u치les son los percentiles 25, 50 y 75 del peso al nacer?

##### `r fa("fas fa-robot", fill = "orange")` Ejemplo 4.8

```{r ejemplo-4_8, exercise = TRUE}
quantile(nac$pesonac, c(0.25, 0.50, 0.75))
```

### 

<font color="#41A9FB">El conjunto de los percentiles 25, 50 y 75 se
denomina cuartiles; dividen el rango de datos en 4 partes iguales, del
25% cada una.

El 25% de los nacidos ten칤an pesos por debajo de los 3096 gramos, el 50%
ten칤an pesos al nacer menores de 3445 gramos y solamente el 25% de los
nacidos ten칤an pesos mayores a 3766 gramos.</font>

### 

#### 4.4 Ejercicios de aplicaci칩n

Ahora es el momento de aplicar los conceptos revisados

Complete las l칤neas de c칩digo en cada una de las celdas, el objetivo es
describir la edad `edadm` de las madres

##### `r fa("fas fa-keyboard", fill = "green")` Ejercicio 4.1

쮺u치l es la edad promedio de las madres en este estudio?

```{r exercise_4_1, exercise = TRUE, exercise.blanks = "___+"}
______(nac$______)
```

```{r exercise_4_1-solution}
mean(nac$edadm)
```

```{r exercise_4_1-check}
grade_code()
```

### 

##### `r fa("fas fa-keyboard", fill = "green")` Ejercicio 4.2

쮺u치l es el rango de edad de las madres?

```{r exercise_4_2, exercise = TRUE, exercise.blanks = "___+"}
______(______$edadm)
```

```{r exercise_4_2-solution}
range(nac$edadm)
```

```{r exercise_4_2-check}
grade_code()
```

### 

##### `r fa("fas fa-keyboard", fill = "green")` Ejercicio 4.3

쮺u치l es la desviaci칩n est치ndar de la edad de las madres?

```{r exercise_4_3, exercise = TRUE, exercise.blanks = "___+"}
______(______)
```

```{r exercise_4_3-solution}
sd(nac$edadm)
```

```{r exercise_4_3-check}
grade_code()
```

### 

##### `r fa("fas fa-keyboard", fill = "green")` Ejercicio 4.4

쮺u치l es el coeficiente de variaci칩n de la edad de las madres?

```{r exercise_4_4, exercise = TRUE, exercise.blanks = "___+"}
______(nac$edadm) / ______(nac$edadm) * ______
```

```{r exercise_4_4-solution}
sd(nac$edadm) / mean(nac$edadm) * 100
```

```{r exercise_4_4-check}
grade_code()
```

### 

##### `r fa("fas fa-keyboard", fill = "green")` Ejercicio 4.5

쯇or debajo de cu치l valor de edad se encuentran el 60% de las madres m치s
j칩venes?

```{r exercise_4_5, exercise = TRUE, exercise.blanks = "___+"}
______(nac$edadm, ______)
```

```{r exercise_4_5-solution}
quantile(nac$edadm, 0.6)
```

```{r exercise_4_5-check}
grade_code()
```

### 

##### `r fa("fas fa-keyboard", fill = "green")` Ejercicio 4.6

Emplee el script encontardo a continuaci칩n, el objetivo es describir el
peso ganado de las madres durante el embarazo `pesogan`. Obtenga las
medidas de tendencia central, dispersi칩n y posici칩n para responder las
siguientes preguntas.

```{r exercise_4_6, exercise = TRUE}

```

```{r quiz_4_1}
quiz(
  question("쮺u치l es el promedio del peso ganado?",
    answer("25 libras"),
    answer("40 libras"),
    answer("32.3 libras", correct = TRUE),
    answer("31 libras"),
    answer("30 libras"),
    allow_retry = TRUE,
    random_answer_order = TRUE
  ),
  question("쮺u치l es la mediana del peso ganado?",
    answer("25 libras"),
    answer("40 libras"),
    answer("32.3 libras"),
    answer("31 libras", correct = TRUE),
    answer("30 libras"),
    allow_retry = TRUE,
    random_answer_order = TRUE
  ),
  question("쮺u치l es la moda del peso ganado?",
    answer("25 libras"),
    answer("40 libras"),
    answer("32.3 libras"),
    answer("31 libras"),
    answer("30 libras", correct = TRUE),
    allow_retry = TRUE,
    random_answer_order = TRUE
  )
)
```

## 5. Gr치ficos

### 5.1 Gr치fico de barras simples

Un gr치fico de barras simples es un gr치fico que representa las categor칤as
de una variable cualitativa con barras rectangulares que tienen una
altura proporcionales a las frecuencias que representa.

El comando para producir un gr치fico de barras simples es `barplot()`, el
comando debe ser aplicado sobre la tabla de frecuencias de la variable
`table()`

##### `r fa("fas fa-robot", fill = "orange")` Ejemplo 5.1

El objetivo es representar gr치ficamente la variable grupo 칠tnico

```{r ejemplo-5_1, exercise = TRUE}
barplot(table(nac$etnia))
```

### 

<font color="#41A9FB"> El gr치fico anterior muestra las frecuencias
absolutas, se puede modificar para representar las frecuencias
porcentuales.</font>

##### `r fa("fas fa-robot", fill = "orange")` Ejemplo 5.2

Se deben calcular las frecuencias relativas, recuerde el comando
`prop.table()`

```{r ejemplo-5_2, exercise = TRUE}
barplot(prop.table(table(nac$etnia))* 100)
```

### 

<font color="#41A9FB">Estos gr치ficos pueden ser mejorados visualmente
empleando algunos par치metros de la funci칩n barplot()</font>

*Par치metros*

-   main="t칤tulo del gr치fico"

-   col = "nombre del color"

-   xlab = "r칩tulo del eje X"

-   ylab = "r칩tulo del eje Y"

-   ylim =c(l칤mite inferior, l칤mite superior)


##### `r fa("fas fa-robot", fill = "orange")` Ejemplo 5.3

```{r ejemplo-5_3, exercise = TRUE}
barplot(prop.table(table(nac$etnia))*100, main = "Figura 2. Porcentaje de madres clasificadas por grupo 칠tnico", col="blue",xlab="Grupo 칠tnico", ylab = "% de madres", ylim=c(0,100))
```

### 

##### `r fa("fas fa-keyboard", fill = "green")` Ejercicio 5.1

Escriba el c칩digo en el script con el objetivo de construir un
gr치fico de barras simples para representar la variable n칰mero previo de
hijos `paridad`

```{r exercise_5_1, exercise = TRUE, exercise.blanks = "___+"}

```

```{r exercise_5_1-solution}
barplot(prop.table(table(nac$paridad))*100 , main = "Figura 3. Porcentaje reci칠n nacidos por n칰mero de hijos en la familia", col="red", xlab = "N칰mero previo de hijos" ,  ylab = "% reci칠n nacidos", ylim = c(0,60))
```

### 5.2 Histograma de frecuencias

Un histograma de frecuencias es un gr치fico para variables cuantitativas
que representa la distribuci칩n de la variable empleando intervalos de
igual amplitud, es una aproximaci칩n a la funci칩n de distribuci칩n de
probabilidad de una variable continua.

El comando para producir un histograma de frecuencias es `hist()`

##### `r fa("fas fa-robot", fill = "orange")` Ejemplo 5.4

El objetivo es representar gr치ficamente la variable edad de la madre

```{r ejemplo-5_4, exercise = TRUE}
hist(nac$edadm)
```

### 

<font color="#41A9FB">El gr치fico pueden visualmente mejorar empleando
los par치metros anteriores</font>

##### `r fa("fas fa-robot", fill = "orange")` Ejemplo 5.5

El objetivo es representar gr치ficamente la distribuci칩n de la variable edad de la madre

```{r ejemplo-5_5, exercise = TRUE}
hist(nac$edadm, main = "Figura 4. Histograma de la edad de la madre", col="yellow", xlab="edad", ylab="N칰mero de madres")
```

### 

##### `r fa("fas fa-keyboard", fill = "green")` Ejercicio 5.2

Escriba el c칩digo en el script con el objetivo de construir un
histograma de frecuencias para la variable a침os de educaci칩n de la madre
`educacion`

```{r exercise_5_2, exercise = TRUE}

```

```{r exercise_5_2-solution}
hist(nac$educacion, main = "Figura 5. Histograma para a침os de educaci칩n de la madre", col="brown", xlab="A침os de educaci칩n", ylab="N칰mero de madres", xlim = c(0,20))
```

### 5.3 Gr치fico de boxplot (caja con bigotes)

Un box plot es un gr치fico para variables cuantitativas que presenta
cinco medidas de resumen num칠rico para la variable: m칤nimo, percentil
25, mediana (percentil 50), percentil 75 y m치ximo.

El comando para producir un histograma de frecuencias es `boxplot()`

##### `r fa("fas fa-robot", fill = "orange")` Ejemplo 5.6

El objetivo es presentar un gr치fico de caja con bigotes para la variable
edad de la madre

```{r ejemplo-5_6, exercise = TRUE}
boxplot(nac$edadm)
```

### 

##### `r fa("fas fa-robot", fill = "orange")` Ejemplo 5.7

Una versi칩n mejorada del gr치fico

```{r ejemplo-5_7, exercise = TRUE}
boxplot(nac$edadm, main="Figura 6. Distribuci칩n de la edad de la madre", col="brown", ylab="Edad")
```

### 

##### `r fa("fas fa-keyboard", fill = "green")` Ejercicio 5.3

Escriba el c칩digo en el script con el objetivo de construir un boxplot
para la variable peso ganado en el embarazo `pesogan`

```{r exercise_5_3, exercise = TRUE}

```

```{r exercise_5_3-solution}
boxplot(nac$pesogan, main = "Figura 6. Distribuci칩n del peso ganado en el embarazo", col="gray", ylab="Edad (a침os)")
```

### 5.4 Gr치fico de dispersi칩n

Un gr치fico de dispersi칩n o *scatterplot* es un gr치fico para visualizar
la relaci칩n entre dos variables cuantitativas, cada una de las variables
se ubica en uno de los ejes del gr치fico.

El comando para producir un histograma de frecuencias es `plot()`

##### `r fa("fas fa-robot", fill = "orange")` Ejemplo 5.8

El objetivo es representar gr치ficamente la relaci칩n entre las variables
peso al nacer y el peso ganado en el embarazo

```{r ejemplo-5_8, exercise = TRUE}
plot(nac$pesonac, nac$pesogan)
```

### 

##### `r fa("fas fa-robot", fill = "orange")` Ejemplo 5.9

Una versi칩n mejorada del gr치fico

```{r ejemplo-5_9, exercise = TRUE}
plot(nac$pesonac, nac$pesogan, main="Figura 8. Relaci칩n entre peso al nacer y peso ganado en el embarazo", pch=19, xlab="Peso al nacer (gramos)",ylab = "Peso ganado (libras) ",col="blue",cex=0.1)
```

### 

##### `r fa("fas fa-keyboard", fill = "green")` Ejercicio 5.4

Escriba el c칩digo en el script con el objetivo de presentar la relaci칩n
entre la edad de la madre y el peso al nacer

```{r exercise_5_4, exercise = TRUE}

```

```{r exercise_5_4-solution}
plot(nac$edadm, nac$pesonac, main="Figura 9. Edad de la madre versus peso al nacer", pch=2, xlab="Edad madre (a침os)",ylab = "Peso al nacer (gramos)",col="black",cex=0.5)

```

### 游 춰El taller ha finalizado! 游녨

Ahora puede encontrar algunos recursos adicionales si su inter칠s es
profundizar en los temas tratados \###

## 6. Recursos adicionales

Si desea profundizar en los t칩picos de estad칤stica descriptiva, el
lenguaje R y los gr치ficos estad칤sticos puede consultar el siguiente
material:

-   [Estad칤stica b치sica](https://www.uv.es/ayala/docencia/nmr/nmr13.pdf)
-   [Estad칤stica descriptiva con
    R](https://bookdown.org/gboccardo/manual-ED-UCH/estadistica-descriptiva-con-rstudio.html)
-   [Gr치ficos estad칤sticos con
    R](https://bookdown.org/jboscomendoza/r-principiantes4/graficas.html)
